Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
# Please add your name and University of Michagan uniqname here...

NAME = 'Liwen Huang'
UMICH_UNIQNAME = 'liwenh'
```

---

```python
version = "v3.3.012821"
```

# SIADS 515 Week 1 Homework (HW1)


## A note about the `%%bash` magic command

The lecture material covered integrating Linux shell commands with Jupyter notebooks using the !, !!, %sx, or %system magic commands. There's one more integration to introduce to you: the %%bash magic command. bash is a popular linux shell, and starting a Jupyter notebook cell with the %%bash magic command will tell Jupyter that the remainder of the cell is to be interpreted as linux (bash) commands.


## A note about the `tee` shell command

For the shell questions in this assignment, you'll generally need to redirect your output to a specifc file for the autograder cells to work. The `tee` command will come in handy, it takes input from STDIN and writes it to STDOUT *and* a file that you specify. This is helpful because you'll be able to see your answer and the file will be writen to allow the autograder to do its work.

```python
%%bash

# This cell cleans up the directory to avoid problems with the autograder.


echo "Cleaning up, removing created files ..."
echo ""
# remove all files in this directory that don't match these patterns
find "." \
    -not -path "." \
    -not -path "./assets" \
    -not -path "./assets/ride.csv" \
    -not -path "./.ipynb_checkpoints" \
    -not -path "./.ipynb_checkpoints/*" \
    -not -path "./assignment1.ipynb" \
    -print -delete
echo ""
echo "... done."
```

Output:
```
Couldn't find program: 'bash'
```

## Example Question

Copy the first 3 lines of `assets/ride.csv` into a file named `solution_example.txt`

*(Note that this is an example problem to demonstrate how to answer shell-based questions in this assignment.)*

```python
%%bash

head -3 assets/ride.csv | tee solution_example.txt
```

<!-- grade_id: cell-1d6b28764fdbf76b, points: 0, grade: true, solution: false -->
```python
# Example Question, Check 1
# 0 points

with open('solution_example.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 3, 'solution file must be 3 lines long'

# There are no hidden autograder tests in this cell.
```

## Question 1
Copy the last 5 lines of `assets/ride.csv` into a file called `solution01.txt`.

<!-- grade_id: cell-d2f755630d46555d, points: 0, grade: false, solution: true -->
```python
%%bash

tail -5 assets/ride.csv | tee solution01.txt
# YOUR CODE HERE


```

<!-- grade_id: cell-1e37b44d1c1726b5, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 1
# 1 point

with open('solution01.txt') as ans:
    lines = ans.readlines()
    
assert len(lines) == 5, 'solution file must be 5 lines long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-dd61bcafaa99e567, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 2
# 1 point

# There are hidden autograder tests in this cell.
```

## Question 2
Copy lines 15 through 20 from `assets/ride.csv` into a file called `solution02.txt`.

<!-- grade_id: cell-018ef04758dd9372, points: 0, grade: false, solution: true -->
```python
%%bash

cut -d ',' -f 15-20 assets/ride.csv | tee solution02.txt
# YOUR CODE HERE


```

<!-- grade_id: cell-450c38d1d8739c1b, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 1
# 1 point

with open('solution02.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 6, 'solution file must be 6 lines long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-0e662ecb874254bf, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 2
# 1 points

# There are hidden autograder tests in this cell.
```

## Question 3
Copy all lines from `assets/ride.csv` that contain "Data,10", "Data,11", "Data,12", or "Data,13" into a file called `solution03.txt`. The lines in the solution should be in the same order as they appear in the ride.csv file.

<!-- grade_id: cell-da32244325922c68, points: 0, grade: false, solution: true -->
```python
%%bash

grep "Data,10|Data,11|Data,12|Data,13" | tee solution03.txt
# YOUR CODE HERE


```

<!-- grade_id: cell-5d6d674650024aa9, points: 1, grade: true, solution: false -->
```python
# Question 3, Check 1
# 1 point

with open('solution03.txt') as ans:
    lines = ans.readlines()

assert len(lines) > 1, 'solution file must have more than one line'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-9d55e0dd789cd561, points: 1, grade: true, solution: false -->
```python
# Question 3, Check 2
# 1 point

# There are hidden autograder tests in this cell.
```

### Question 4

Count the lines from `assets/ride.csv` that contain "Data,9". Create a file called `solution04.txt` that has one line that contains the number of lines in ride.csv. That line should not contain anything other than the number of lines from 

<!-- grade_id: cell-0cf2e81e963c5402, points: 0, grade: false, solution: true -->
```python
%%bash

# YOUR CODE HERE


```

<!-- grade_id: cell-e175ee39f82440ef, points: 1, grade: true, solution: false -->
```python
# Question 4, Check 1
# 1 point

with open('solution04.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 1, 'solution file should be 1 line long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-036ce5a0a73e018b, points: 1, grade: true, solution: false -->
```python
# Question 4, Check 2
# 1 point

# There are hidden autograder tests in this cell.

# Checking if line 1 of solution04.txt ONLY contains the correct answer.
```

## Three solutions to the same problem...

For Question 5, 6, 7, you will create a CSV file for each with the same content based on the data from `assets/ride.csv`. You'll use a different solution for each problem:
  - Question 5: with shell commands
  - Question 6: with Python (without Pandas)
  - Question 7: with Pandas


You'll notice that although the `ride.csv` file is a csv file, it looks a bit different than most csv files.  The first few columns are always of the form (using regular expression notation) ```"(Data|Definition),[\d]+,record"```.  After those columns, the data appear as a key-value-unit triple.  For example:

```
Data,9,record,timestamp,"896018545",s,position_lat,"504719750",semicircles,position_long,"-998493490",semicircles,distance,"10.87",m,altitude,"285.79999999999995",m,speed,"1.773",m/s,unknown,"3929",,unknown,"1002",,enhanced_altitude,"285.79999999999995",m,enhanced_speed,"1.773",m/s,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,
```

In this case, you have:

```
timestamp,"896018545",s
position_lat,"504719750",semicircles
position_long,"-998493490",semicircles
speed,"1.773",m/s
```


The data we are interested in comes from the lines in `assets/ride.csv` that start with "Data,9,record". 

For each question, you will create a CSV file with these headers, and this first file:

| timestamp| latitude| longitude| distance|          altitude| speed|
|----------|---------|----------|---------|------------------|------|
| 896018545|504719750|-998493490|    10.87|285.79999999999995| 1.773|
|       ...|      ...|       ...|      ...|               ...|   ...|
|       ...|      ...|       ...|      ...|               ...|   ...|

## Question 5 - Shell

Create a new file that matches the format described above.
  - The file should be named `solution05.csv`
  - Solve the problem using only Linux shell commands
  
**For this more complicated problem, the `tee` command might be more trouble than it's worth.**

<!-- grade_id: cell-3ddf21c4bdc70e67, points: 0, grade: false, solution: true -->
```python
%%bash

# YOUR CODE HERE


```

<!-- grade_id: cell-7652086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 5, Check 1
# 1 point

import csv

def close_match(a, b):
    return round(float(a), 3) == round(float(b), 3)

def ride_csv_check1(filename):
    with open(filename) as ans:
        ans_reader = csv.reader(ans, delimiter=',', quotechar='"')

        header = next(ans_reader)
        assert header[0] == 'timestamp', 'incorrect header'
        assert header[1] == 'latitude', 'incorrect header'
        assert header[2] == 'longitude', 'incorrect header'
        assert header[3] == 'distance', 'incorrect header'
        assert header[4] == 'altitude', 'incorrect header'
        assert header[5] == 'speed', 'incorrect header'

        first_line = next(ans_reader)
        # checks values to three decimal places
        assert close_match(first_line[0], 896018545), 'incorrect value'
        assert close_match(first_line[1], 504719750), 'incorrect value'
        assert close_match(first_line[2], -998493490), 'incorrect value'
        assert close_match(first_line[3], 10.87), 'incorrect value'
        assert close_match(first_line[4], 285.79999999999995), 'incorrect value'
        assert close_match(first_line[5], 1.773), 'incorrect value'

        
solution_filename = 'solution05.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-4f3a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 5, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-c179efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 5, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```

## Q06 - Python without Pandas

Create a new file that matches the format described above:
  - The file should be named `solution06.csv`
  - Don't use any shell or Jupyter magic commands
  - Solve the problem using Python
  - You may use the Python standard library (e.g. the built-in [csv](https://docs.python.org/3/library/csv.html) package)
  - You may NOT use third-party libraries (e.g. Pandas)

<!-- grade_id: cell-e53a157e0c081aba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


```

<!-- grade_id: cell-q652086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 6, Check 1
# 1 point

# note that this function is defined in the cell for Question 5, Check 1 above
solution_filename = 'solution06.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-q63a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 6, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-q679efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 6, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```

## Q07 -  Pandas

Create a new file that matches the format described above:
  - The file should be named `solution07.csv`
  - Don't use any shell or Jupyter magic commands
  - Solve the problem using functionality provided by the Pandas library in Python

<!-- grade_id: cell-184c965f0270227f, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# YOUR CODE HERE


```

<!-- grade_id: cell-q752086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 7, Check 1
# 1 point

# note that this function is defined in the cell for Question 5, Check 1 above
solution_filename = 'solution07.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-q73a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 7, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-q779efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 7, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```